Longitudinal Awake Imaging of Mouse Deep Brain Microvasculature with Super-resolution Ultrasound Localization Microscopy

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public Review):

Summary:

Wang and Colleagues present a study aimed at demonstrating the feasibility of repeated ultrasound localization microscopy (ULM) recording sessions on mice chronically implanted with a cranial window transparent to US. They provided quantitative information on their protocol, such as the required number of Contrast enhancing microbubbles (MBs) to get a clear image of the vasculature of a brain coronal section. Also, they quantified the co-registration quality over time-distant sessions and the vasodilator effect of isoflurane.

Strengths:

The study showed a remarkable performance in recording precisely the same brain coronal section over repeated imaging sessions. In addition, it sheds light on the vasodilator effect of isoflurane (an anesthetic whose effects are not fully understood) on the different brain vasculature compartments, although, as the Authors stated, some insights in this aspect have already been published with other imaging techniques. The experimental setting and protocol are very well described.

Wang and co-authors submitted a revised version of their study, which shows improvements in the clarity of the data description.

However, the flaws and limitations of this study are substantially unchanged.

The main issues are:

Statistics are still inadequate. The TOST test proposed in this revised version is not equivalent to an ANOVA. Indeed, multivariate analyses should be the most appropriate, given that some quantifications were probably made on multiple vessels from different mice. The 3 reviewers mentioned the flaws in statistics as the primary concern.

Response 01: We thank the reviewer for raising this important point. We fully acknowledge the limitations of our current statistical analysis. We would like to clarify that the TOST procedure was applied exclusively to the measurements taken from the same vessel segment in the same animal across different time points, with the purpose of evaluating the consistency of vessel diameter measurements. We recognize that the statistical analysis in this study remains limited, which we have acknowledged as a key limitation in the manuscript. This constraint arises primarily from the limited number of animals, and our analysis should be interpreted as a representative case study rather than a generalized statistical conclusion. We have revised the manuscript to clarify these points and to more explicitly acknowledge the statistical limitations.

(Line 329) “Our current study primarily focused on demonstrating the feasibility of longitudinal ULM imaging in awake animals, instead of conducting a systematic investigation of how isoflurane anesthesia alters cerebral blood flow. Due to the limited number of animals used, the analyses presented in this work should be interpreted as example case studies. While the trends observed across animals were consistent, the small sample size restricts the scope of statistical inference. For future work, it would be valuable to design more rigorous control experiments with larger sample sizes to systematically compare the effects of isoflurane anesthesia, awake states, and other anesthetics that do not induce vasodilation on cerebral blood flow.”

No new data has been added, such as testing other anesthetics.

Response 02: We acknowledge that the current study does not include data involving other anesthetics, and we have also discussed this point in our initial response. In fact, we did attempt to use other anesthetics such as ketamine. However, we found it difficult to draw reliable conclusions due to experimental limitations such as variable anesthesia recovery profiles and injection timing, as elaborated in the following paragraphs. Therefore, we decided not to include these data in the current study to avoid potential misinterpretation.

One major limitation of our experimental setup is that imaging in the awake state is necessarily conducted after a brief period of isoflurane-anesthesia. This brief anesthesia allows for the intravenous injection of microbubbles via the tail vein. Isoflurane is particularly suited for this purpose due to its rapid onset and offset. Mice can recover quickly once the gas is withdrawn, which enables relatively consistent post-anesthesia imaging in the awake state.

In contrast, other anesthetic agents present challenges. Their recovery profiles are slower, more variable, and less controllable. Reversal drugs can be administered to awaken the animals, but they add another variability. These may lead to greater fluctuations in cerebral hemodynamics and factors introduce uncertainty in the timing of bolus microbubble injection. As such, our current setup is not ideal for systematically comparing different anesthetics and could yield misleading results.

A more appropriate strategy for comparing awake ULM imaging with different anesthetics would be performing awake imaging first, followed by imaging under anesthesia. This would ensure that the awake condition is free from residual anesthetic effects. However, this method raises higher requirement in bubble delivery, as no anesthesia can be used for the intravenous injection.

To address this, we are actively exploring another solution using indwelling jugular vein catheterization. By surgically implanting a catheter into the jugular vein prior to imaging, we can establish a stable and reproducible route for microbubble delivery in fully awake animals without any anesthesia induction. This method has the potential to enable direct and reliable comparisons across different physiological states. However, the implementation of this technique and the associated experimental findings go beyond the scope of the current study and will be presented in a future manuscript.

In the present work, we have emphasized the methodological limitations of our approach and clarified that our primary goal is to highlight the necessity and feasibility of awake-state ULM imaging. The focus is not to comprehensively characterize the effects of different anesthetic agents on microvascular brain flow. We appreciate your understanding and interest in this important future direction.

Based the responses and previous revision, we have further refined the discussion of the relevant limitations:

(Line 324) “Although isoflurane is widely used in ultrasound imaging because it provides long-lasting and stable anesthetic effects, it is important to note that the vasodilation observed with isoflurane is not representative of all anesthetics. Some anesthesia protocols, such as ketamine combined with medetomidine, do not produce significant vasodilation and are therefore preferred in experiments where vascular stability is essential, such as functional ultrasound imaging. Our current study primarily focused on demonstrating the feasibility of longitudinal ULM imaging in awake animals, instead of conducting a systematic investigation of how isoflurane anesthesia alters cerebral blood flow. Due to the limited number of animals used, the analyses presented in this work should be interpreted as example case studies. While the trends observed across animals were consistent, the small sample size restricts the scope of statistical inference. For future work, it would be valuable to design more rigorous control experiments with larger sample sizes to systematically compare the effects of isoflurane anesthesia, awake states, and other anesthetics that do not induce vasodilation on cerebral blood flow.”

(Line 347) “Another limitation of this study is the potential residual vasodilatory effect of isoflurane anesthesia on awake imaging sessions and the short imaging window available after bolus injection. The awake imaging sessions were conducted shortly after the mice had emerged from isoflurane anesthesia, required for the MB bolus injections. The lasting vasodilatory effects of isoflurane may have influenced vascular responses, potentially contributing to an underestimation of differences in vascular dynamics between anesthetized and awake state. In addition, since microbubbles are rapidly cleared from circulation, the duration of effective imaging is limited to only a few minutes, which also overlaps with the anesthesia recovery period, constraining the usable awake-state imaging window. Future improvement on microbubble infusion using an indwelling jugular vein catheter presents a promising alternative to address these limitations. This method allows for stable microbubble infusion without the need for anesthesia induction, ensuring that the awake imaging condition is free from residual anesthetic effects. Moreover, it has the potential to extend the duration of imaging sessions, offering a longer and more stable time window for data acquisition. Furthermore, by performing ULM imaging in the awake state first, instead of starting with anesthetized imaging, researchers can achieve a more rigorous comparison of how various anesthetics influence cerebral microvascular dynamics relative to the awake baseline.”

The Authors still insist on using the term Vascularity which they define as: 'proportion of the pixel count occupied by blood vessels within each ROI, obtained by binarizing the ULM vessel density maps and calculating the percentage of the pixels with MB signal.'. Why not use apparent cerebral blood volume or just CBV? Introducing an unnecessary and redundant term is not scientifically acceptable. In this revised version, vascularity is also used to indicate a higher vascular density (Line 275), which does not make sense: blood vessels do not generate from the isoflurane to the awake condition in a few minutes. Rev2 also raised this point.

Response 03: Thank you for revisiting this important point. We acknowledge that the term vascularity is difficult to interpret for readers, and we also recognize that we did not sufficiently justify its use in the earlier version.

Based on your suggestion, we have now replaced all instances of “vascularity” with “fractional vessel area”. While the underlying definition remains the same, fractional vessel area offers a more intuitive description. The term “fractional” denotes that the vessel area is normalized to the total area of the selected ROI. This normalization is essential for fair comparisons across ROIs of different sizes, such as Figures 4i–k to evaluate various brain regions. We would also like to clarify that this was not introduced as an unnecessary or redundant term, but rather as a more suitable metric for longitudinal ULM analysis. We did consider using apparent cerebral blood volume (CBV), estimated from microbubble counts. However, we found that it was less robust and meaningful in the context of longitudinal ULM comparisons. Below we provide further justification for using the vessel area instead:

(1) Using the vessel area is more robust:

In longitudinal ULM comparisons, normalization across time points is essential to enable fair and meaningful comparisons. In our study, we normalized the data based on a cumulative 5 million microbubbles (e.g., Fig. 2). Other normalization strategies could also be adopted, as long as the resulting vascular maps reach a sufficiently saturated state. However, even with normalization, it remains important to use a quantitative metric that is minimally biased and invariant to experimental fluctuations across time points. Vessel area, derived from binarized vessel maps, is less sensitive to variations in acquisition time and microbubble concentration. This is because repeated microbubble trajectories through the same location are not counted multiple times. In contrast, apparent CBV, calculated from the microbubble counts, is more susceptible to different concentration conditions. Since repeated detections in the same location accumulate, the metric can be dependent on injection efficiency and imaging duration. While CBV may still be valid under well-controlled, steady-state conditions, we found the vessel area to be a more robust and reliable metric for longitudinal analysis under our current bolus-injection protocol.

(2) Using the vessel area is more meaningful:

Compared to CBV, the vessel area provides a more direct representation of structural characteristics such as vessel diameter. Anesthesia-induced vasodilation leads to an increase in vessel diameter. Although local diameter changes can be assessed by manually selecting vessel segments, this approach is labor-intensive and prone to selection bias. To enable a more comprehensive and objective assessment of such morphological changes, fractional vessel area provides a more informative alternative to CBV, as it captures diameter-related variations at a global or regional scale, and avoids potential biases associated with manually selecting specific vessels or regions.

In response to: vascularity is also used to indicate a higher vascular density (Line 275), which does not make sense: blood vessels do not generate from the isoflurane to the awake condition in a few minutes.

We agree that blood vessels cannot be generated in a few minutes. Vascularity (now fractional vessel area) should be interpreted as apparent vessel density, which reflects a probabilistic estimate of vessel density based on the detectable microbubble.

Both apparent vessel density and apparent CBV are indirect, sampling-based approximations of vascular features, and both are fundamentally limited by microbubble detection sensitivity. Low microbubble concentrations lead to underestimation of both CBV and vessel area. A change from zero to non-zero in these metrics does not imply the physical appearance or disappearance of vessels, but rather reflects a change in the likelihood of detecting flow in each region.

In summary, while neither fractional vessel area (vascularity in previous versions) nor apparent CBV is a perfect metric due to the inherent limitations of ULM, we believe the vessel area provides a more robust and meaningful parameter for our longitudinal comparisons. We have revised the main text to include this explanation and acknowledge the limitations and interpretation of fractional vessel area more explicitly.

Revision in Results:

(Line 181) “To validate the broader applicability of our findings, we conducted ROI-based analyses using fractional vessel area and mean velocity as primary metrics. These metrics extended the analysis of vessel diameter and flow velocity to entire brain regions or selected ROIs, which provides a more objective assessment of cerebral blood flow changes at a global scale and reduces the bias associated with manually selecting vessel segments. For vessel area measurements, the term fractional denotes that the vessel area is normalized to the total area of the selected ROI. This normalization is essential for fair comparisons across ROIs of different sizes.”

Revision in Methods: definition of vascularity

(Line 571) “In ROI-based analysis, we focused on two primary parameters: fractional vessel area and mean velocity. Fractional vessel area was defined as the proportion of the pixel count occupied by blood vessels within each ROI, obtained by binarizing the ULM vessel density maps and calculating the percentage of the pixels with MB signal. Mean velocity was calculated by averaging all non-zero pixel of velocity estimates within the ROI. The velocity distribution within each ROI was also visualized using violin plots, as shown in Fig. 2, 4 and 6, to illustrate the range and density of flow velocity estimates across different acquisition. In this study, we focused on these two metrics because they represent the most straightforward extension of single-vessel analysis to brain-wide vascular changes.”

We put our ROI analysis code on GitHub and added a “Code availability” section. We hope it can serve as a foundation for users to explore different quantitative metrics in their own longitudinal ULM studies. We hope to provide an example to inspire further exploration.

(Line 578) “Code availability

To support quantitative longitudinal analysis of ULM data, we developed an open-source MATLAB application (https://github.com/ekerwang/ULMQuantitativeAnalysis). This tool is designed to facilitate ROI-based analysis of ULM images for longitudinal comparisons. It supports multiple quantification metrics, including but not limited to vessel area and mean velocity used in this study. Users can select and adapt different metrics based on their specific applications, as a wide range of ULM-based quantification metrics have been developed for different pathological and pharmacological studies.”

The long-term recordings mentioned by the Authors refer to the 3-week time frame analyzed in this study. However, within each acquisition, the time available from imaging is only a few minutes (< 10', referring to most of the plots showing time courses) after the animals' arousal from isoflurane and before bubbles disappear. This limitation should be acknowledged.

Response 04: Thank you for this comment. We agree that the current imaging sessions are constrained by the short time window available after the animal’s arousal from isoflurane and before bubbles disappear. This limitation indeed restricts the duration of usable awake-state imaging in our current bolus injection protocol. As discussed earlier, we are actively exploring the use of a jugular vein catheterization approach to address this limitation. This approach has the potential to extend the imaging session duration and provide a longer, more stable time window. We have now acknowledged this limitation more explicitly in the revised Discussion section.

(Line 347) “Another limitation of this study is the potential residual vasodilatory effect of isoflurane anesthesia on awake imaging sessions and the short imaging window available after bolus injection. The awake imaging sessions were conducted shortly after the mice had emerged from isoflurane anesthesia, required for the MB bolus injections. The lasting vasodilatory effects of isoflurane may have influenced vascular responses, potentially contributing to an underestimation of differences in vascular dynamics between anesthetized and awake state. In addition, since microbubbles are rapidly cleared from circulation, the duration of effective imaging is limited to only a few minutes, which also overlaps with the anesthesia recovery period, constraining the usable awake-state imaging window. Future improvement on microbubble infusion using an indwelling jugular vein catheter presents a promising alternative to address these limitations. This method allows for stable microbubble infusion without the need for anesthesia induction, ensuring that the awake imaging condition is free from residual anesthetic effects. Moreover, it has the potential to extend the duration of imaging sessions, offering a longer and more stable time window for data acquisition. Furthermore, by performing ULM imaging in the awake state first, instead of starting with anesthetized imaging, researchers can achieve a more rigorous comparison of how various anesthetics influence cerebral microvascular dynamics relative to the awake baseline.”

The more precise description of the number of mice and blood vessels analyzed in Figure 6 makes it apparent the limited number of independent samples used to support the findings of this work. A limitation that should be acknowledged. The newly provided information added as Supplementary Figure 1 should be moved to the main text, eventually in the figure legends. The limited data in support of the findings was also highlighted by Rev2 and, indirectly, by Rev3.

Response 05: We acknowledge the limited number of independent samples used in this study. In the revised manuscript, we have explicitly emphasized this limitation in the Discussion section. Specifically, we added the following statement:

(Line 329) “Our current study primarily focused on demonstrating the feasibility of longitudinal ULM imaging in awake animals, instead of conducting a systematic investigation of how isoflurane anesthesia alters cerebral blood flow. Due to the limited number of animals used, the analyses presented in this work should be interpreted as example case studies. While the trends observed across animals were consistent, the small sample size restricts the scope of statistical inference. For future work, it would be valuable to design more rigorous control experiments with larger sample sizes to systematically compare the effects of isoflurane anesthesia, awake states, and other anesthetics that do not induce vasodilation on cerebral blood flow.”

Following your suggestion, we have also moved the newly provided information (the table in Supplementary Figure 1) into figure captions. In addition, we have modified in the Methods section to ensure that this information is clear.

(Line 406) “Eight healthy female C57 mice (8-12 weeks) were used for this study, numbered as Mouse 1 to Mouse 8. Three mice (Mouse 1–3) were used to compare imaging results between awake and anesthetized states (Fig. 3 and 4). Three additional mice (Mouse 4–6) underwent longitudinal imaging over a three-week period (Fig. 5 and 6). Among them, Mouse 4 was also used as an example to demonstrate the overall system schematic and saturation conditions (Fig. 1 and 2). Several mice (Mouse 2, 6, 7, and 8) exhibited suboptimal cranial window quality or image artifacts and were included to illustrate common surgical or imaging issues (Supplementary Fig. 1). The specific usage of each animal is also annotated in the corresponding figure captions.”

Reviewer #2 (Public Review):

The authors present a very interesting collection of methods and results using brain ultrasound localization microscopy (ULM) in awake mice. They emphasize the effect of the level of anesthesia on the quantifiable elements assessable with this technique (i.e. vessel diameter, flow speed, in veins and arteries, area perfused, in capillaries) and demonstrate the possibility of achieving longitudinal cerebrovascular assessment in one animal during several weeks with their protocol.

The authors made a good rewriting of the article based on the reviewers' comments. One of the message of the first version of the manuscript was that variability in measurements (vessel diameter, flow velocity, vascularity) were much more pronounced under changes of anesthesia than when considering longitudinal imaging across several weeks. This message is now not quite mitigated, as longitudinal imaging seems to show a certain variability close to the order of magnitude observed under anesthesia. In that sense, the review process was useful in avoiding hasty conclusion and calls for further caution in ULM awake longitudinal imaging, in particular regarding precision of positioning and cancellation of tissue motion.

Strengths:

Even if the methods elements considered separately are not new (brain ULM in rodents, setup for longitudinal awake imaging similar to those used in fUS imaging, quantification of vessel diameters/bubble flow/vessel area), when masterfully combined as it is done in this paper, they answer two questions that have been longrunning in the community: what is the impact of anesthesia on the parameters measured by ULM (and indirectly in fUS and other techniques)? Is it possible to achieve ULM in awake rodents for longitudinal imaging? The manuscript is well constructed, well written, and graphics are appealing.

The manuscript has been much strengthened by the round of review, with more animals for the longitudinal imaging study.

Weaknesses:

Some weaknesses remain, not hindering the quality of the work, that the authors might want to answer or explain.

When considering fig 4e and fig 4j together: it seems that in fig 4e the vascularity reduction in the cortical ROI is around 30% for downward flow, and around 55% for upward flow; but when grouping both cortical flows in fig 4j, the reduction is much smaller (~5%), even at the individual level (only mouse 1 is used in fig 4e). Can you comment on that?

Response 06: Thank you for carefully pointing this out. This discrepancy arises primarily from differences in ROI selections.

The vascularity metric (now we changed the term into fractional vessel area, based on Reviewer 1’s comments) is calculated as the proportion of vessel-occupied pixels relative to the total ROI area. As such, it is best suited for longitudinal comparisons within the same ROI rather than across-ROI comparisons, particularly when the size and vessel composition of the ROIs differ.

In Fig. 4e, the cortical ROI includes mostly the penetrating vessels, which are selected due to their clear distinction between upward (venous) and downward (arterial) flow directions. Pial vessels were intentionally excluded because flow direction alone does not reliably distinguish arteries from veins in these surface vessels. Thus, the goal of this analysis was to indicate arteriovenous differences, rather than to represent the full cortical vascular changes.

In contrast, the ROIs used in Fig. 4j aim to provide a more comprehensive view of cortical vascular responses without distinguishing flow direction. That’s why both penetrating and pial vessels are included. Since pial vessels showed relatively smaller vascularity changes within the coronal cross-sections analyzed in our study, their inclusion in the cortical ROI likely contributed to the smaller overall reduction in vascularity observed in Figure 4j.

To address this potential confusion, we have added further clarification in the Results section of the revised manuscript.

(Line 209) “It is worth noting that prior analyses (Fig. 4d–h) aimed to illustrate arteriovenous differences. Since pial vessels are difficult to distinguish as arteries or veins based on flow direction in coronal plane imaging, they were excluded from the ROI selection in those analyses. In the current whole-brain comparisons (Fig. 4i-k), the cortical ROIs no longer exclude pial vessels, since distinguishing between arteries and veins is not required. This aims to provide a more comprehensive representation of cortical vasculature.”

When considering fig 4e, fig 4j, fig 6e and fig 6i altogether, it seems that vascularity can be highly variable, whether it be under anesthesia or vascular imaging, with changes between 5 to 40%. Is this vascularity quantification worth it (namely, reliable for example to quantify changes in a pathological model requiring longitudinal imaging)?

Response 07: Thank you for raising this important point. We found that imaging in the awake state is inherently more variable than under anesthesia. In contrast, anesthetized imaging offers a more controlled and stable physiological condition, as anesthesia suppresses many sources of variation. For pathological studies, if the vascular or hemodynamic changes induced by anesthesia do not interfere with the scientific question being addressed, imaging under anesthesia can still be a practical and effective approach, due to its experimental simplicity and better physiological consistency.

The higher variability observed in awake imaging arises from both physiological fluctuations in animals and unavoidable experimental inconsistencies, such as small misalignment on the imaging plane across sessions. If the research question aims to avoid the confounding effects of anesthesia, then instead of suppressing variation through anesthesia, it is important to acknowledge the natural baseline variation in the awake state. However, efforts should be made to minimize technical sources of variation. We have added a brief discussion of this issue at the end of the manuscript to reflect this consideration.

(Line 396) “However, it is also important to note that although longitudinal awake imaging presents promise to avoid the confounding effects of anesthetics, imaging under anesthesia remains more convenient and controllable in many cases. For applications where the physiological question of interest is not sensitive to anesthesia-induced vascular effects, anesthetized imaging still offers a simpler and more stable approach. Awake imaging inherently exhibits greater physiological variability. However, care must be taken at the experimental level to minimize confounding sources of variation, such as stress level of the animal or handling inconsistencies, to ensure that the measurements are physiologically meaningful.”

Regarding whether fractional vessel area (formerly referred to as vascularity) is a worthwhile metric for longitudinal quantification: based on our experience and comparisons, we found vessel area to be relatively robust and informative (see also Response 02 to Reviewer 1 for details). However, we acknowledge that other quantitative metrics—such as microbubble count, tortuosity, or flow directionality—may be more suitable depending on the specific pathological model or research question. How these metrics perform in awake imaging and longitudinal disease models is indeed an open and important question. We hope our work can serve as a foundation to inspire further investigation in this direction. To facilitate such exploration, we have developed and open-sourced a MATLAB-based analysis tool that supports multiple quantitative ULM metrics for longitudinal comparison. We encourage users to adapt and extend this framework to evaluate different quantitative metrics.

(Line 578) “Code availability

To support quantitative longitudinal analysis of ULM data, we developed an open-source MATLAB application (https://github.com/ekerwang/ULMQuantitativeAnalysis). This tool is designed to facilitate ROI-based analysis of ULM images for longitudinal comparisons. It supports multiple quantification metrics, including but not limited to vessel area and mean velocity used in this study. Users can select and adapt different metrics based on their specific applications, as a wide range of ULM-based quantification metrics have been developed for different pathological and pharmacological studies.”

Reviewer #2 (Recommendations For The Authors):

Images in figure 4 lack color bars.

Response 08: Thank you for pointing this out. The color bars for the images in Figure 4 are the same as those used in the corresponding images in Figure 3. We have now added the explanation of color bars to the revised version of Figure 4 caption.

Fig 4d: upward and downward are probably swapped.

Response 09: Thank you for pointing this out, and we apologize for the oversight. They were mistakenly swapped. We have corrected this error in the revised figure.

No quantitative conclusions are drawn regarding the changes in vessel diameter under anesthesia? Is it not significant? If it is not then why bring changes in diameter to our attention in fig 3 (white arrows) and figure 4b?

Response 10: Our intention in highlighting diameter changes in Figure 3 (white arrows) and Figure 4b was to provide an illustrative example of isoflurane-induced diameter changes at the single-vessel level. These examples are meant to serve as case studies, not as the basis for broad statistical conclusions.

In the initial version of the manuscript, we attempted to draw quantitative conclusions by measuring vessel diameters from ten manually selected vessel segments at each location. However, based on feedback from other reviewers, we decided to remove this analysis in the revised version. Manual selection of vessel segments is highly subjective and prone to bias, limiting its reliability for quantitative interpretation.

Instead, we focused on ROI-based analysis using fractional vessel area (formerly referred to as vascularity), which reflects widespread changes in vessel diameter across regions. It is a more generalizable and less biased metric for quantifying vascular diameter changes.

We further explained this in the Results section:

(Line 181) “To validate the broader applicability of our findings, we conducted ROI-based analyses using fractional vessel area and mean velocity as primary metrics. These metrics extended the analysis of vessel diameter and flow velocity to entire brain regions or selected ROIs, which provides a more objective assessment of cerebral blood flow changes at a global scale and reduces the bias associated with manually selecting vessel segments. For vessel area measurements, the term fractional denotes that the vessel area is normalized to the total area of the selected ROI. This normalization is essential for fair comparisons across ROIs of different sizes.”

Line 210 "In summary, statistical analysis revealed a decrease in individual vessel diameter" this does not seem to be supported by this version of the manuscript as no analysis is done on a representative group of vessels for the diameter.

Response 11: Thank you for pointing out this important issue. In line with our previous response (Response 10), we would like to clarify that the analysis of individual vessel diameter was intended to serve as an example study, rather than a statistically supported conclusion based on a group of vessels. To avoid confusion, we have removed the phrase “statistical analysis revealed a decrease in individual vessel diameter” from the manuscript.

The meaning of the *** in fig 6b and 6c should be clarified as: -it is not explicitly stated - the equivalence test interpretation is less usual than other tests.

Response 12: We thank the reviewer for pointing out this important issue. We agree that the use of asterisks (***) in Fig. 6b and 6c may have led to confusion, as such markers are typically associated with statistical significance in difference testing. In our case, the analysis was based on the two one-sided test (TOST) procedure to assess statistical equivalence, which is indeed less commonly used and could be misinterpreted.

To address this, we have replaced the asterisks *** in the figure with the label “equiv.”, which more clearly reflects the intended interpretation. Additionally, we have revised the figure caption and the main text to explicitly state that these markers denote statistical equivalence (not difference) as determined by TOST, with the equivalence margin defined as three times the standard deviation of one week.

(Figure 6 Caption) “Statistical analysis was performed using the two one-sided test (TOST) to evaluate consistency of measurement. The label “equiv.” indicates statistically equivalent measurements (p < 0.001), defined as interweek differences smaller than three times the standard deviation of one week.”

(Line 240) “Statistical testing of equivalence was conducted using the two one-sided test (TOST) procedure, which evaluates whether the difference between two time points falls within a predefined equivalence margin. Specifically, equivalence is defined as the inter-week difference being smaller than three times the standard deviation of one week. A statistically significant result in TOST (p < 0.001) supports the interpretation that the measurements are statistically equivalent, which is denoted as “equiv.” in the figures.”

Line 237 and following: please consider rephrasing into "To further generalize these findings and examine longitudinal variation in ROI-based analysis, we used Mouse 4 as an example to show the consistency of blood flow density across different flow directions in the cortex (Fig. 6d) and extended the quantitative analysis to all three mice (Fig. 6e) (individual ULM upward and downward flow images for all three mice over the threeweek longitudinal study period can be found in Supplementary Fig. 4)." The paragraph will make much more sense.

Response 13: We appreciate your helpful rephrasing. We have fully adopted your proposed revision to enhance the clarity and coherence of the text. The sentence now reads exactly as you recommended:

(Line 250): “To further generalize these findings and examine longitudinal variation in ROI-based analysis, we used Mouse 4 as an example to show the consistency of blood flow density across different flow directions in the cortex (Fig. 6d) and extended the quantitative analysis to all three mice (Fig. 6e) (individual ULM upward and downward flow images for all three mice over the three-week longitudinal study period can be found in Supplementary Fig. 4).”

Line 248: "While arterial and venous flow velocity distributions exhibit clear distinctions, their variations over the three weeks remained acceptable" the meaning of acceptable remains elusive.

Response 14: Thank you for pointing out the ambiguity in the phrase “remained acceptable”. To improve clarity and precision, we have revised the sentence to provide a more informative description. The updated sentence now reads:

(Line 261) “While arterial and venous flow velocity distributions exhibit clear distinctions, the distribution shapes remained relatively consistent across the three weeks. Specifically, variation in median velocity were within 1 mm/s. In contrast, anesthesia-induced changes can lead to velocity shifts exceeding 1 mm/s.”

Line 253: consider rephrasing in "Despite subcortical regions showing the largest vascularity variability consecutive to anesthesia-induced changes, vascularity in those regions was relatively stable values in the longitudinal study" as otherwise the link between the 2 parts of the sentence feels odd.

Response 15: Thank you for your constructive suggestion regarding the logical flow of the sentence. We fully agree with your point and have revised the sentence exactly as you proposed.

(Line 268) “Despite subcortical regions showing the largest vascularity variability consecutive to anesthesia-induced changes, vascularity in those regions was relatively stable values in the longitudinal study.”

Changes from the earlier version to the revised version

Abstract

Abstract

Ultrasound localization microscopy (ULM) is an emerging imaging modality that resolves microvasculature in deep tissues with high spatial resolution. However, existing preclinical ULM applications are largely constrained to anesthetized animals, introducing confounding vascular effects such as vasodilation and altered hemodynamics. As such, ULM quantifications (e.g., vessel diameter, density, and flow velocity) may be confounded by the use of anesthesia, undermining the usefulness of ULM in practice. Here we introduce a method to address this limitation and achieve ULM imaging in awake mouse brain. Pupillary monitoring was used to support the presence of the awake state during ULM imaging. Vasodilation induced by isoflurane was observed by ULM. Upon recovery to the awake state, reductions in vessel density and flow velocity were observed across different brain regions. In the cortex, the effects induced by isoflurane are more pronounced on venous flow than on arterial flow. In addition, serial in vivo imaging of the same animal brain at weekly intervals demonstrated the highly robust longitudinal imaging capability of the proposed technique. The consistency was further verified through quantitative analysis on individual vessels, cortical regions of arteries and veins, and subcortical regions. This study demonstrates longitudinal ULM imaging in the awake mouse brain, which is crucial for many ULM brain applications that require awake and behaving animals.

Introduction

Sensitive imaging of correlates of activity in the awake brain is fundamental for advancing our understanding of neural function and neurological diseases. In the field of neuroscience, there is a growing interest in utilizing imaging techniques to study the rodent brain, which serves as a valuable model for investigating brain function1, 2. Brain imaging modalities for rodents encompass a wide range of techniques, including but not limited to functional magnetic resonance imaging (fMRI)3, 4, positron emission tomography (PET)5, 6, one- and two-photon imaging7, 8, photoacoustic imaging9, 10, and more recently ultrasound localization microscopy (ULM)11. ULM is uniquely capable of imaging microvasculature situated in deep tissue (e.g., at a depth of several centimeters). ULM can also be combined with the principles of functional ultrasound (fUS)12–15 to image whole-brain neural activity at a microscopic scale16. The structural and functional imaging capabilities of ULM have opened new doors for numerous basic research and clinical applications that involve cerebral microvasculature17–19.

At present, a key limitation associated with existing ULM brain imaging studies is the use of anesthesia, which induces profound alterations to cerebral blood flow (CBF), including changes in vessel size (e.g., diameter) and flow velocity20. As such, CBF measurements under anesthesia do not reflect the blood flow under the normal physiological state of the brain. In addition, anesthetics also have a significant attenuating effect on neural responses to sensory inputs, thereby impacting the neurovascular coupling process21–26. Therefore, in neuroscience research, brain imaging in the awake state is often preferred over imaging under anesthesia (e.g., fluorescence imaging27, photoacoustic imaging28, 29, fMRI4, PET6, and fUS30–, 33). As ULM is gaining traction in many preclinical brain imaging applications, enabling ULM for awake animals has become crucial to eliminate the confounding vascular effects of anesthetics and obtain accurate structural and functional cerebrovascular measurements.

Another challenge associated with preclinical ULM brain imaging is to conduct long-term, longitudinal studies, which are essential for tracking disease progression or therapeutic impacts for many neurological disease applications34–36. The key technical challenge for longitudinal ULM brain imaging is to find the same imaging plane and reconstruct consistent ULM images across different imaging sessions. Misalignment of imaging planes or tissue movement will undermine ULM imaging quality and result in inconsistent cerebrovascular quantifications. Although intact skull imaging has been shown to be feasible for short-term studies (e.g., over a few days)18, long-term monitoring can be challenging because changes in skull properties over time (e.g., thickness and composition) could negatively impact ULM imaging quality. Currently there is a strong need for methodological developments to enable longitudinal brain imaging with ULM.

In this study, we developed a method for awake and longitudinal ULM brain imaging in a mouse model to eliminate the confounding vascular effects of anesthesia and enable monitoring of cerebral vasculature over a three-week period in the same animal. We constructed a head-fixed awake imaging platform and established a ULM image reconstruction metric to allow comparisons of ULM images acquired at different states of wakefulness (e.g., awake vs. anesthesia). Our method allowed detailed comparisons of local and global variations of the cerebral vasculature and blood flow under awake and anesthesia conditions. Detailed quantitative analysis of vessel diameter and blood flow velocity was performed. We also demonstrated robust longitudinal ULM brain imaging on same animals with high repeatability over multiple weeks.

Results

ULM brain imaging can be performed in head-fixed awake mice

To achieve consistent ULM brain imaging while allowing limited movement in awake animals, a head-fixed imaging platform with a chronic cranial window was used in this study. Pre-surgery handling and post-surgery habituation were employed to alleviate animal stress and facilitate awake imaging33. To further minimize tissue motion, a 3D-printed body tube was adopted33 for animal immobilization (Fig. 1a). A headpost was implanted during cranial window surgery to fix the animal’s head to the body tube. In this procedure, the skull was removed, and the brain tissue was protected using a polymethyl pentene (PMP) membrane37, which was further protected by a layer of silicone rubber (Fig. 1b). Animals were imaged after one week of surgical recovery and head-fixation habituation. Prior to each imaging session, the silicone rubber was gently removed with forceps, and the transparent PMP membrane allowed detailed examination of the brain surface to ensure absence of tissue damage (Fig. 1c).

Experimental setup for awake mice ULM brain imaging.

a, 3D model of the body tube enabling rapid fixation of the headpost onto the tube. b, Top view photograph of the mouse after cranial window surgery, with the headpost protected by silicone rubber. c, Photograph taken after the post-surgery recovery, with the removal of the silicone rubber protection. d, Image captured by the camera positioned in front of the mouse during the imaging session. e, Change of the pupillary area over time, after microbubble injection (T = 0s), along with magnified portion of the pupil photo in d at six time points (T=0, 200, 400, 600, 800, and 1000 second). Pupil is outlined with red dashed circle. f, Awake ULM image obtained from data collected during a fully awake state, as indicated by the gray shaded area in e. (scale bar: 1 mm)

Contrast enhancing microbubbles (MBs) were administered intravenously via tail vein in this study (details provided in Methods). Isoflurane anesthesia was terminated right after microbubble (MB) injection, marking the beginning of the data acquisition (i.e., T = 0s). Mice were allowed to gradually regain consciousness. Throughout the awakening process, an infrared camera was used to monitor the pupil to provide a reference for the state of arousal38 (Fig. 1d). Fig. 1e shows the change in pupillary area over time, which reveals a gradual enlargement of the pupil after the cessation of anesthesia. Due to differences in tail vein injection timing and anesthesia depth, the time required for each mouse to fully awaken varied. Although it was not feasible to get pupil size stabilized just after 500 seconds for each animal, ULM reconstruction only used the data that acquired after the animal reached full pupillary dilation, to ensure that ULM accurately captures the cerebrovascular characteristics in the awake state. Ultrafast plane wave imaging data were collected for ULM reconstruction (details provided in Methods) and a representative awake ULM brain image is presented in Fig. 1f.

Microbubble count serves as a quantitative metric for awake ULM image reconstruction

Due to the stochastic nature of MB localization and fluctuations of MB concentration in the blood stream, it is challenging to compare ULM images acquired at different states of wakefulness with different cerebral blood flow conditions. Therefore, to obtain complete, fully saturated ULM images under different physiological conditions and MB concentrations, we used MB count as a measure of vessel saturation to determine the completeness of ULM reconstructions. To ensure high quality ULM imaging, only MBs that were successfully tracked by the uTrack algorithm (details provided in Methods) were considered as effective MB signals that were utilized for ULM image quantification.

To facilitate equitable comparison of brain perfusion at different states, a practical saturation point enabling stable quantification of most vessels needs to be established. Our observations indicated that when the cumulative MB count reached 5 million, ULM images achieved a relatively stable state. Accordingly, in this study, the saturation point was defined as a cumulative MB count of 5 million. There are also possible alternatives for ULM image normalization. For example, different ULM images can be normalized to have the same saturation rate. However, the proposed method of using the same number of cumulative MB count for normalization enables the analysis of blood flow distribution across different brain regions from a probabilistic perspective. The following analysis substantiates this criterion.

Fig. 2a compares ULM directional vessel density maps and flow speed maps generated with 1, 3, 5, and 6 million MBs, using the same animal as shown in Fig. 1. To quantitatively confirm saturation, multiple vessel segments were selected for further analysis. Fig. 2b presents the measured vessel diameter for a specific segment at various MB counts. After binarizing the ULM map, the vessel diameter was measured by calculating the distance from the vessel centerline to the edge. Each point along the centerline of the segment provided a diameter measurement, enabling calculation of the mean and standard error. At low MB counts, vessels appeared incompletely filled, leading to inaccurate estimation of vessel diameter due to incomplete profiles. For example, at 1–2 million MBs, the binarized ULM map displayed a width of only one or two pixels along the segment. As a result, the measurements always yielded the same diameter values (two pixels, ∼10um) with a consistently low standard error of the mean across the entire segment. With increased MB counts, the measured vessel diameter gradually rose, ultimately reaching saturation. The plots in Fig. 2b show that vessel diameter stabilized at 5 million MB count. Additionally, Fig. 2c illustrates the changes in flow velocity measured at different cumulative MB counts. The violin plots display the distribution of flow speed estimates for all valid centerline pixels within the selected segment. At low MB counts (1–3 million), flow velocity estimates fluctuated, but they stabilized as the MB count increased (4–6 million MBs). At 5 million MBs, flow velocity estimates were nearly identical to those at 6 million MBs, corroborating previous findings that vessel velocity measurements stabilize as MB count grows39. To assess the generalizability of the 5 million MB saturation condition, vessel segments from three different mice across various brain regions were examined. The results, shown in Supplementary Fig. 2, confirm that this saturation criterion applies broadly. Although the 5 million MB threshold may not ensure absolute saturation for all vessels, it is generally effective for vessels larger than 15 μm. This MB count threshold was therefore adopted as a practical criterion.

Data processing standards for awake mice ULM imaging.

a, ULM directional vessel density maps and flow velocity maps at cumulative MB count of 1, 3, 5, and 6 million. b, Vessel diameter measurements along the selected segment at varying cumulative MB counts. Each point along the vessel centerline provides a measurement, with the mean and standard error of the mean (SEM) plotted to show average diameter and variability. c, Flow velocity measurements along the selected segment at varying cumulative MB counts. Each non-zero centerline pixel provides a velocity value, with the distribution across positions shown as a violin plot. d, Time courses of MB count in each second (blue curve) and the cumulative MB count (orange curve) starting from T = 500s. The vertical gray dashed lines indicate the time points when the cumulated MB count reaches 1, 2, 3, 4, 5, and 6 million. e, Time courses of saturation level of the ULM image (blue curve) and the filling rate of pixels (orange curve). The filling rate is calculated by taking the derivative of the filled pixel count, and then normalized to the initial filling rate at the beginning of ULM reconstruction (T = 500s). f, Relationship between pixel filling rate and cumulative MB count, eliminating the time axis by plotting the orange curves from d and e together. This figure presents a case study based on the same mouse shown in Fig 1. The x-axis for d-f begins at 500 seconds because, at this point, the mouse’s pupil size stabilized, indicating it had recovered to an awake state. Consequently, ULM images were accumulated starting from this time. It is important to note that not every mouse requires 500 seconds to fully awaken; the time to reach a stable awake state varies across individual mice.

To further verify that the proposed MB bolus injection method can help to achieve ULM image saturation shortly after mice awaken from anesthesia, an analysis on the change in MB concentration over time was conducted once pupil size had stabilized (T = 500s). Fig. 2d shows a clear trend of decreasing MB concentration in the blood stream and increasing cumulative MB count with time. Fig. 2e demonstrates a flattened vessel saturation curve and a rapidly reducing vessel filling rate, which is typical for ULM reconstruction39. The pixel filling rate dropped below 5% of the initial rate after 300 seconds of ULM data accumulation (T = 800s), indicating ULM image saturation. Fig. 2f further examines the relationship between the pixel filling rate and the cumulative MB count, which is independent of data acquisition time. The pixel filling rate at 5 million cumulative MBs is always below 5% of the initial rate for each experiment, ensuring image saturation. In summary, all the quantitative measurements indicate that ULM images obtained using the proposed metric (i.e., 5 million MBs) were complete and can be used to consistently measure cerebral vascular properties such as vessel diameter and blood flow velocity. All subsequent ULM images in this study were produced using this criterion.

Earlier version: ULM reveals increase in blood flow induced by isoflurane anesthesia

Revised version: ULM reveals an increase in blood flow induced by isoflurane anesthesia

Earlier version: Fig. 3 presents a comparison of ULM directional vessel density maps and flow velocity maps in awake and anesthetized states for three different coronal planes from three animals. Four regions of interest (ROIs) were selected within each imaging plane to provide detailed comparisons. When comparing vessel density maps, ULM images that were acquired in the awake state demonstrate a global reduction of vascularity, which refers to percentage of pixels that occupied by blood vessels. The reduction is also clearly observed in magnified local regions especially for regions 5, which encompasses the pretectal region. In addition to reduced vascularity, ULM also reveals decreased vessel diameter which reflects vasoconstriction after the animals woke up (e.g., white arrows in regions 1, 3, and 6, corresponding to the thalamus, cortex and midbrain/cortex overlap region). Focusing on ULM flow speed maps, a global reduction in flow speed can be clearly observed when transitioning from anesthetized to awake. Regional maps further revealed the significant flow speed reduction for most of the vessels throughout the brain (e.g., blue arrows).

Revised version: Fig. 3 presents a comparison of ULM directional vessel density maps and flow velocity maps in awake and anesthetized states for three different coronal planes from three animals. Four regions of interest (ROIs) were selected within each imaging plane to provide detailed comparisons. When comparing vessel density maps, ULM images that were acquired in the awake state demonstrate a global reduction of vessel density. The reduction is also clearly observed in magnified local regions especially for regions 5, which encompasses the pretectal region. ULM also reveals decreased vessel diameter which reflects vasoconstriction after the animals woke up (e.g., white arrows in regions 1, 3, and 6, corresponding to the thalamus, cortex and midbrain/cortex overlap region). Focusing on ULM flow speed maps, a global reduction in flow speed can be clearly observed when transitioning from anesthetized to awake. Regional maps further revealed the significant flow speed reduction for most of the vessels throughout the brain (e.g., blue arrows).

Comparison of ULM images in isoflurane anesthetized and awake states.

Earlier version: The ULM images of three different mice are shown. The upper panel shows the comparison of directional vessel density maps, while the lower panel shows the comparison of flow velocity. Four regions of interest (ROIs) are selected within each coronal plane (indicated by white dashed boxes in the whole-plane view of the vessel density map) for zoom-in comparison.

Revised version: The ULM images of three different mice (Mouse 1-3) are shown. The upper panel shows the comparison of directional vessel density maps, while the lower panel shows the comparison of flow velocity. Four regions of interest (ROIs) are selected within each coronal plane (indicated by white dashed boxes in the whole-plane view of the vessel density map) for zoom-in comparison.

To quantitatively compare the results from Fig. 3, we performed various analyses in Fig. 4. Since previous studies have shown that isoflurane induces different effects on arteries and veins20, 40, 41, we first selected a cortical artery and vein for comparison. In Fig. 4a, a cortical region near Region 3 from Fig. 3 was selected, and one artery and one vein in this region were analyzed. Fig. 4b compares the vessel diameters of these segments in anesthetized and awake states. For the selected artery, no significant change in diameter was observed, with both anesthetized and awake states showing an average diameter of 22 µm. In contrast, the selected vein displayed a notable response to isoflurane, with its diameter decreasing significantly from 54 µm under isoflurane anesthesia to 41 µm in the awake state. Flow velocity analysis of the same vessels further suggests that isoflurane may have a greater impact on venous blood flow velocity: arterial flow velocity increased slightly from 7.33 mm/s under anesthesia to 8.04 mm/s in the awake state, whereas venous flow velocity decreased from 10.48 mm/s to 8.11 mm/s, with both the magnitude and statistical significance of the decrease greater in the vein.

Quantitative comparison of the ULM images in isoflurane anesthetized and awake states.

Earlier version: a, Example cortical region near Region 3 in Mouse 1 (from Fig. 3), with vessel density maps displayed separately by flow direction. In this cortical region, upward flow corresponds to venous blood, while downward flow corresponds to arterial blood. One artery and one vein segment were selected within the boxed areas for analysis of flow velocity and diameter. b,c, Comparisons of vessel diameter (b) and flow velocity (c) for the selected arterial and venous segments. Statistical analysis was conducted using t-test at each measurement point along the segments. d, Whole cortical region in Mouse 1, showing separate blood flow density maps for upward (venous) and downward (arterial) flow, with a defined ROI (white circled region) used for quantitative analysis. e,f, Quantitative assessment of vascularity (e) and mean flow velocity (f) within the cortical ROI for Mouse 1–3, presented as individual values in each state (left y-axis) and the relative reduction from the anesthetized baseline to the awake state (right y-axis). g, Flow velocity map for the same cortical region shown in d, using Mouse 1 as an example. h, Flow velocity distributions within the cortical ROI for Mouse 1–3, comparing anesthetized and awake states. i, Regional analysis across different brain areas using bidirectional ULM maps, with selected ROIs from the cortex (CTX), hippocampal formation (HPF), thalamus (Thal), and midbrain (MBr) for each mouse. j,k, Quantitative analysis of vascularity (j) and mean blood flow velocity (k) for the ROIs in each brain region in Mouse 1–3. (*: p<0.05, **: p<0.01, ***: p<0.001)

Revised version: a, Example cortical region near Region 3 in Mouse 1 (from Fig. 3), with vessel density maps displayed separately by flow direction. In this cortical region, upward flow corresponds to venous blood, while downward flow corresponds to arterial blood. One artery and one vein segment were selected within the boxed areas for analysis of flow velocity and diameter. b,c, Comparisons of vessel diameter (b) and flow velocity (c) for the selected arterial and venous segments. Statistical analysis was conducted using t-test at each measurement point along the segments. d, Whole cortical region in Mouse 1, showing separate blood flow density maps for upward (venous) and downward (arterial) flow, with a defined ROI (white circled region) used for quantitative analysis. e,f, Quantitative assessment of fractional vessel area (e) and mean flow velocity (f) within the cortical ROI for Mouse 1–3, presented as individual values in each state (left y-axis) and the relative reduction from the anesthetized baseline to the awake state (right y-axis). g, Flow velocity map for the same cortical region shown in d, using Mouse 1 as an example. h, Flow velocity distributions within the cortical ROI for Mouse 1–3, comparing anesthetized and awake states. i, Regional analysis across different brain areas using bidirectional ULM maps, with selected ROIs from the cortex (CTX), hippocampal formation (HPF), thalamus (Thal), and midbrain (MBr) for each mouse. j,k, Quantitative analysis of fractional vessel area (j) and mean blood flow velocity (k) for the ROIs in each brain region in Mouse 1–3. Colorbars of a, d, g are the same as in Fig. 3. (*: p<0.05, **: p<0.01, ***: p<0.001)

Added in the revised version: While single-vessel analysis provides valuable insights, its generalizability is limited. To validate the broader applicability of our findings, we conducted ROI-based analyses using fractional vessel area and mean velocity as primary metrics. These metrics extended the analysis of vessel diameter and flow velocity to entire brain regions or selected ROIs, which provides a more objective assessment of cerebral blood flow changes at a global scale and reduces the bias associated with manually selecting vessel segments. For vessel area measurements, the term fractional denotes that the vessel area is normalized to the total area of the selected ROI. This normalization is essential for fair comparisons across ROIs of different sizes.

Earlier version: While single-vessel analysis provides valuable insights, its generalizability is limited. To validate the broader applicability of our findings, we conducted ROI-based analyses of cortical arteries and veins, using vascularity and mean velocity as primary metrics. The relationship between arteriovenous classification and blood flow direction is more straightforward to infer on cortex, with downward flow representing cortical arteries and upward flow representing cortical veins16. Fig. 4d shows, for Mouse 1, the comparison of cortical upward (venous) and downward (arterial) flow density maps under different states. Supplementary Fig. 3 also provides a detailed comparison of upward and downward flow for all three mice in both anesthetized and awake conditions. Fig. 4e demonstrates that transitioning from isoflurane anesthesia to the awake state led to a reduction in vascularity for both arteries and veins within the selected cortical ROI across all three mice. However, the decrease in venous vascularity (averaging 55% across mice) was greater than that of arterial vascularity (averaging 35%). Fig. 4f further illustrates the change in mean velocity within the selected ROI, where both arteries and veins showed a decreasing trend; however, venous flow velocity reduced by an average of 38%, compared to a 19% reduction in arterial flow velocity across the cortical ROI. These results underscore the different responses of arteries and veins to isoflurane anesthesia, which is also indicated by previous studies20,40,41.

Revised version: Our ROI-based analysis began with cortical arteries and veins. The relationship between arteriovenous classification and blood flow direction is more straightforward to infer on cortex, with downward flow representing cortical arteries and upward flow representing cortical veins16. Fig. 4d shows, for Mouse 1, the comparison of cortical upward (venous) and downward (arterial) flow density maps under different states. Supplementary Fig. 3 also provides a detailed comparison of upward and downward flow for all three mice in both anesthetized and awake conditions. Fig. 4e demonstrates that transitioning from isoflurane anesthesia to the awake state led to a reduction in vessel area for both arteries and veins within the selected cortical ROI across all three mice. However, the decrease in venous vessel area (averaging 55% across mice) was greater than that of arterial (averaging 35%). Fig. 4f further illustrates the change in mean velocity within the selected ROI, where both arteries and veins showed a decreasing trend; however, venous flow velocity reduced by an average of 38%, compared to a 19% reduction in arterial flow velocity across the cortical ROI. These results underscore the different responses of arteries and veins to isoflurane anesthesia, which is also indicated by previous studies20, 40, 41.

To visualize flow velocity changes more clearly, we highlighted the flow velocity comparison for the same region in Fig. 4g and provided a violin plot of velocity estimates within the ROI in Fig. 4h. The violin plots reveal a consistently greater reduction in venous flow velocity from anesthesia to the awake state across all three mice.

Earlier version: For deeper brain regions, flow analyses were conducted without separating arteries and veins, as this physiological separation is not established for ULM in these regions. Fig. 4i displays the selected ROIs for each mouse, while Fig. 4j and 4k present the trends in vascularity and flow velocity changes across different brain regions. Results indicate that vascularity and flow velocity decreased across the brain as mice transitioned from anesthetized to awake states, with subcortical ROIs showing more pronounced changes than cortical ROIs. For example, in the thalamus and midbrain regions, average vascularity decreased by 37% and average flow velocity by 17%, whereas cortical ROIs exhibited smaller decreases of 4% in vascularity and 11% in flow velocity. This work is consistent with a previous fMRI study, where it was also observed that isoflurane-induced cerebral hyperemia was not most pronounced in the cerebral cortex compared with other deeper brain regions42.

Revised version: For deeper brain regions, flow analyses were conducted without separating arteries and veins, as this physiological separation is not established for ULM in these regions. Fig. 4i displays the selected ROIs for each mouse, while Fig. 4j and 4k present the trends in vessel area and flow velocity changes across different brain regions. It is worth noting that prior analyses (Fig. 4d–h) aimed to illustrate arteriovenous differences. Since pial vessels are difficult to distinguish as arteries or veins based on flow direction in coronal plane imaging, they were excluded from the ROI selection in those analyses. In the current whole-brain comparisons (Fig. 4i-k), the cortical ROIs no longer exclude pial vessels, since distinguishing between arteries and veins is not required. This aims to provide a more comprehensive representation of cortical vasculature. Results indicate that vessel area and flow velocity decreased across the brain as mice transitioned from anesthetized to awake states, with subcortical ROIs showing more pronounced changes than cortical ROIs. For example, in the thalamus and midbrain regions, fractional vessel area decreased by 37% on average and flow velocity decreased by 17%, whereas cortical ROIs exhibited smaller decreases of 4% in fractional vessel area and 11% in flow velocity. This work is consistent with a previous fMRI study, where it was also observed that isoflurane-induced cerebral hyperemia was not most pronounced in the cerebral cortex compared with other deeper brain regions42.

Earlier version: In summary, statistical analysis revealed a decrease in individual vessel diameter, vascularity, and blood flow velocity (particularly in venules) after awakening. These findings align with existing research, indicating higher blood perfusion during isoflurane anesthesia20,40,41,43,44.

Revised version: In summary, statistical analysis revealed a decrease in vessel area and blood flow velocity (particularly in venules) after awakening. These findings align with existing research, indicating higher blood perfusion during isoflurane anesthesia20, 40, 41, 43, 44.

Awake ULM imaging demonstrates high consistency in longitudinal imaging across different weeks

Longitudinal awake ULM brain imaging was feasible using the surgical and imaging techniques presented in this study. Fig. 5 presents the results of awake brain imaging performed on the same brain region over three consecutive weeks. Two ROIs at different depths were selected for each mouse to compare micro vessel reconstruction across different time points. Fig. 5 demonstrates a high level of consistency in both directional vessel density maps and flow speed maps obtained from the three imaging sessions, although some minor discrepancies were observed. The inconsistency could potentially be attributed to physiological variation and/or slight misalignment.

Longitudinal awake ULM imaging results on the same coronal plane for three consecutive weeks.

Earlier version: The ULM images of three different mice are shown. The upper panel shows the comparison of directional vessel density maps, while the lower panel shows the comparison of flow velocity. Two regions of interest (a cortical ROI and a sub cortical ROI) are selected within each coronal plane (indicated by white dashed boxes in the whole-plane view of the vessel density map) for zoom-in comparison.

Revised version: The ULM images of three different mice (Mouse 4-6) are shown. The upper panel shows the comparison of directional vessel density maps, while the lower panel shows the comparison of flow velocity. Two regions of interest (a cortical ROI and a sub cortical ROI) are selected within each coronal plane (indicated by white dashed boxes in the whole-plane view of the vessel density map) for zoom-in comparison.

Earlier version: To quantitatively evaluate the consistency of ULM imaging in longitudinal studies, we conducted a similar analysis to that shown in Fig. 4. First, a case study was performed on individual vessels by selecting an artery and a vein in the cortex of Mouse 4 (Fig. 6a) and measuring vessel diameter (Fig. 6b) and flow velocity (Fig. 6c) along the selected segments. The results show small variation in measurements across weeks. For vessel diameter, both arterial and venous segments demonstrated statistical equivalence across the three weekly measurements. Statistical testing of equivalence was conducted using the two one-sided test (TOST) procedure, which evaluates the null hypothesis that the difference between the two weeks is larger than three times the standard deviation of one week. Large p-values from the TOST analyses indicate rejection of the null hypothesis, supporting that the measurements are equivalent. For flow velocity, apart from a slight decrease in the venous flow during the second week (4.49 mm/s, compared to 6.03 mm/s in the first week and 6.00 mm/s in the third week), no other significant differences were observed. Overall, quantitative measurements maintained high consistency, though slight variations were expected due to differences in probe positioning and physiological variations among the mice.

Revised version: To quantitatively evaluate the consistency of ULM imaging in longitudinal studies, we conducted a similar analysis to that shown in Fig. 4. First, a case study was performed on individual vessels by selecting an artery and a vein in the cortex of Mouse 4 (Fig. 6a) and measuring vessel diameter (Fig. 6b) and flow velocity (Fig. 6c) along the selected segments. The results show small variation in measurements across weeks. For vessel diameter, both arterial and venous segments demonstrated statistical equivalence across the three weekly measurements. Statistical testing of equivalence was conducted using the two one-sided test (TOST) procedure, which evaluates whether the difference between two time points falls within a predefined equivalence margin. Specifically, equivalence is defined as the inter-week difference being smaller than three times the standard deviation of one week. A statistically significant result in TOST (p < 0.001) supports the interpretation that the measurements are statistically equivalent, which is denoted as “equiv.” in the figures. For flow velocity, apart from a slight decrease in the venous flow during the second week (4.49 mm/s, compared to 6.03 mm/s in the first week and 6.00 mm/s in the third week), no other significant differences were observed. Overall, quantitative measurements maintained high consistency, though slight variations were expected due to differences in probe positioning and physiological variations among the mice.

Analysis of ULM images acquired in the three consecutive weeks.

Earlier version: a, Two example cortical regions in Mouse 4 (from Fig. 5), with one for artery selection and the other for vein. The vessel segments were selected within the boxed areas for analysis of flow velocity and diameter. b,c, Comparisons of vessel diameter (b) and flow velocity (c) for the selected arterial and venous segments. Statistical analysis was conducted using the two one-sided test (TOST) procedure, which evaluates the null hypothesis that the difference between the two weeks is larger than three times the standard deviation of one week. d, Whole cortical region in Mouse 4, showing separate blood flow density maps for upward (venous) and downward (arterial) flow, with a defined ROI (white circled region) used for quantitative analysis. e, Bar plot for the vascularity comparison of the cortical artery and vein among Mouse 4-6. f, Flow velocity map for the same cortical region shown in d, using Mouse 4 as an example. g, Flow velocity distributions within the cortical ROI for Mouse 4-6, comparing arterial and venous flow across three weeks. h, Bidirectional ULM vessel density map of subcortical ROIs for the three different mice. i, Bar plot for the vascularity comparison of the subcortical ROIs in h. j, ULM flow velocity map of the same subcortical ROIs for the three different mice. k, Bar plot for the mean velocity comparison of the subcortical ROIs. (*: p<0.05, **: p<0.01, ***: p<0.001)

Revised version: a, Two example cortical regions in Mouse 4 (from Fig. 5), with one for artery selection and the other for vein. The vessel segments were selected within the boxed areas for analysis of flow velocity and diameter. b,c, Comparisons of vessel diameter (b) and flow velocity (c) for the selected arterial and venous segments. Statistical analysis was performed using the two one-sided test (TOST) to evaluate consistency of measurement. The label “equiv.” indicates statistically equivalent measurements (p < 0.001), defined as inter-week differences smaller than three times the standard deviation of one week. d, Whole cortical region in Mouse 4, showing separate blood flow density maps for upward (venous) and downward (arterial) flow, with a defined ROI (white circled region) used for quantitative analysis. e, Bar plot for the fractional vessel area comparison of the cortical artery and vein among Mouse 4-6. f, Flow velocity map for the same cortical region shown in d, using Mouse 4 as an example. g, Flow velocity distributions within the cortical ROI for Mouse 4-6, comparing arterial and venous flow across three weeks. h, Bidirectional ULM vessel density map of subcortical ROIs for the three different mice. i, Bar plot for the fractional vessel area comparison of the subcortical ROIs in h. j, ULM flow velocity map of the same subcortical ROIs for the three different mice. k, Bar plot for the mean velocity comparison of the subcortical ROIs.

Earlier version: To further generalize these findings and examine longitudinal variation in ROI-based analysis, we used Mouse 4 as an example to show the consistency of blood flow density across different flow directions in the cortex (Fig. 6d) and extended the quantitative analysis to all three mice (Fig. 6e). In Supplementary Fig. 4, we present a comparison of upward and downward flow images for all three mice over the three-week longitudinal study period. In the comparative study shown in Fig. 4, clear differences were seen between arterial and venous responses under various conditions, while in the longitudinal study, both arteries and veins exhibited similar levels of variation. For instance, in Mouse 6, both arterial and venous vascularity showed a notable decrease from around 60% in the first week to around 40% in the second week, which is a relatively large variation compared with other measurements. This quantitative analysis aims to inform readers of the expected range of variation of this technique. Fig. 6f displays the flow velocity map for the same region as in Fig. 6d, with flow velocity distributions inside the selected ROI for all three mice shown in Fig. 6g. While arterial and venous flow velocity distributions exhibit clear distinctions, their variations over the three weeks remained acceptable. Even in the previously mentioned case of Mouse 6 during Week 1 and Week 2, arterial mean flow velocity only shifted from 6.31 mm/s in the first week to 7.00 mm/s in the second week.

Revised version: To further generalize these findings and examine longitudinal variation in ROI-based analysis, we used Mouse 4 as an example to show the consistency of blood flow density across different flow directions in the cortex (Fig. 6d) and extended the quantitative analysis to all three mice (Fig. 6e) (individual ULM upward and downward flow images for all three mice over the three-week longitudinal study period can be found in Supplementary Fig. 4). In the comparative study shown in Fig. 4, clear differences were seen between arterial and venous responses under various conditions, while in the longitudinal study, both arteries and veins exhibited similar levels of variation. For instance, in Mouse 6, both arterial and venous fractional vessel area showed a notable decrease from around 60% in the first week to around 40% in the second week, which is a relatively large variation compared with other measurements. This quantitative analysis aims to inform readers of the expected range of variation of this technique. Fig. 6f displays the flow velocity map for the same region as in Fig. 6d, with flow velocity distributions inside the selected ROI for all three mice shown in Fig. 6g. While arterial and venous flow velocity distributions exhibit clear distinctions, the distribution shapes remained relatively consistent across the three weeks. Specifically, variation in median velocity were within 1 mm/s. In contrast, anesthesia-induced changes can lead to velocity shifts exceeding 1 mm/s. Even in the previously mentioned case of Mouse 6 during Week 1 and Week 2, arterial mean flow velocity only shifted from 6.31 mm/s in the first week to 7.00 mm/s in the second week.

Earlier version: We also conducted ROI analyses without distinguishing between arteries and veins for subcortical regions (Fig. 6h–k), using the same ROIs as those shown in Fig. 5. Although subcortical regions show larger difference to anesthesia-induced changes, they maintained relatively stable values in the longitudinal study. In terms of vascularity (Fig. 6i), Mouse 5 exhibited the greatest variation among the three mice, with vascularity increasing from 55% in the first week to 67% in the third week. Compared to the maximum 40% reduction seen from anesthesia to awake states in Fig. 4j, vascularity measurement in subcortical regions remained consistent. This observation also applied to mean velocity measurements.

Revised version: We also conducted ROI analyses without distinguishing between arteries and veins for subcortical regions (Fig. 6h–k), using the same ROIs as those shown in Fig. 5. Despite subcortical regions showing the largest vessel area variability in anesthesia-induced changes, vessel area in those regions was relatively stable in the longitudinal study. In terms of vessel area (Fig. 6i), Mouse 5 exhibited the greatest variation among the three mice, with vessel area increasing from 55% in the first week to 67% in the third week. Compared to the maximum 40% reduction seen from anesthesia to awake states in Fig. 4j, vessel area measurement in subcortical regions remained consistent. This observation also applied to mean velocity measurements.

Discussion

In this study, we introduced a method for performing ULM brain imaging in awake mice under a longitudinal study setting. Our method enabled high-resolution imaging of deep cerebral micro-vasculature with the animal under the awake state. We translated the awake imaging techniques previously described in fUS33 to our study to enable awake ULM, and established a quantitative metric for ULM image reconstruction. Based on the setup above, we studied CBF changes induced by anesthesia, which aligned well with literature. Isoflurane has been shown to increase vascular diameter and CBF in mice, as validated by multiple imaging modalities including optical coherence tomography40, photoacoustic microscopy41, two-photon microscopy43, and laser speckle imaging20, 44. These effects have also been validated in larger animal models such as rats42, dogs45, and marmosets46. In human, vasodilation and increased CBF caused by volatile anesthetics such as isoflurane have also been reported22.

Earlier version: Statistical analysis from Fig. 4 shows that certain vessels exhibit a larger diameter under isoflurane anesthesia, and the vascularity, calculated as the percentage of vascular area within selected brain region ROIs, is also higher in the anesthetized state. These findings suggest a vasodilation effect induced by isoflurane, consistent with existing research20,40,41,43,44. It is worth noting that although our data indicate a global elevation of CBF under isoflurane anesthesia, individual vessels exhibit large discrepancies in behavior. For example, the vessel at the left lower corner from Mouse 3 in the entorhinal cortex (Fig. 3) shows almost no blood flow during anesthesia but then exhibits high vessel pixel density after awakening. The wide range of vessel behaviors were also previously reported in literature20,46. Our results indicate that awake ULM imaging has ample spatial resolution and imaging depth of penetration to resolve individual vessel variations down to micron-sized vessels deep into the brain. This is a unique capability that is not available from other biomedical imaging modalities.

Revised version: Statistical analysis from Fig. 4 shows that certain vessels exhibit a larger diameter under isoflurane anesthesia, and the fractional vessel area, calculated as the percentage of vascular area within selected brain region ROIs, is also higher in the anesthetized state. These findings suggest a vasodilation effect induced by isoflurane, consistent with existing research20, 40, 41, 43, 44. It is worth noting that although our data indicate a global elevation of CBF under isoflurane anesthesia, individual vessels exhibit large discrepancies in behavior. For example, the vessel at the left lower corner from Mouse 3 in the entorhinal cortex (Fig. 3) shows almost no blood flow during anesthesia but then exhibits high vessel pixel density after awakening. The wide range of vessel behaviors were also previously reported in literature20, 46. Our results indicate that awake ULM imaging has ample spatial resolution and imaging depth of penetration to resolve individual vessel variations down to micron-sized vessels deep into the brain. This is a unique capability that is not available from other biomedical imaging modalities.

Increased blood flow velocity induced by isoflurane has also been reported by other studies40, 41. However, previous research presented different speculations on the predominant factor contributing to the increase in CBF induced by anesthesia, specifically whether the increase is attributed to vasodilation or increase in blood flow velocity20. One study found significant changes in both blood flow and vessel diameter but minor changes in flow velocity, suggesting that the increase in blood flow was largely driven by vasodilation40. Conversely, another study drew the opposite conclusion20. Benefitting from the large field of view (FOV) of ULM and its capability to directly quantify microvascular blood flow velocity, we can make a more comprehensive inference regarding the relationship among the changes in vessel diameter, flow speed, and flow volume from anesthetized to awake states. For arteries, the change in blood flow velocity is not significant, indicating that the alteration in blood flow may be primarily due to vasodilation instead of velocity change. Isoflurane causes vasodilation by acting on the ion channels (e.g., potassium channel) of smooth muscle45, which is more abundantly found in arteries than in veins. In the case of veins, which do not actively dilate or constrict, their vessel diameter and blood flow variations are more likely controlled by passive mechanisms. Fig. 5 reveals significant differences in flow velocity of veins between anesthesia and awake state, suggesting that the changes in flow velocity may have a greater impact on venous blood flow volume compared with arterial volume.

The differences in cerebral vasculature between anesthetized and awake states observed using ULM are also in agreement with other studies20, 40, 41, 44. However, previous studies mostly used optical imaging techniques, which have limited penetration depth and can only observe surface pial vessels in the cortex. Some other studies using fMRI can detect deeper CBF changes in the whole brain, but they do not provide insights about small vessel blood flow variations due to insufficient spatial resolution42. As a bridging imaging modality between MRI and optical techniques, awake ULM enables observations of detailed microvascular variations induced by anesthesia across the whole depth of the brain, which provides complementary information to existing biomedical imaging modalities.

Earlier version: Although isoflurane is widely used in ultrasound imaging because it provides long-lasting and stable anesthetic effects, it is important to note that the vasodilation observed with isoflurane is not representative of all anesthetics. Some anesthesia protocols, such as ketamine combined with medetomidine, do not produce significant vasodilation and are therefore preferred in experiments where vascular stability is essential, such as functional ultrasound imaging47. Therefore, in future studies, it would be valuable to design more rigorous control experiments with larger sample sizes to systematically compare the effects of isoflurane anesthesia, awake states, and other anesthetics that do not induce vasodilation on cerebral blood flow.

Revised version: Although isoflurane is widely used in ultrasound imaging because it provides long-lasting and stable anesthetic effects, it is important to note that the vasodilation observed with isoflurane is not representative of all anesthetics. Some anesthesia protocols, such as ketamine combined with medetomidine, do not produce significant vasodilation and are therefore preferred in experiments where vascular stability is essential, such as functional ultrasound imaging47. Our current study primarily focused on demonstrating the feasibility of longitudinal ULM imaging in awake animals, instead of conducting a systematic investigation of how isoflurane anesthesia alters cerebral blood flow. Due to the limited number of animals used, the analyses presented in this work should be interpreted as example case studies. While the trends observed across animals were consistent, the small sample size restricts the scope of statistical inference. For future work, it would be valuable to design more rigorous control experiments with larger sample sizes to systematically compare the effects of isoflurane anesthesia, awake states, and other anesthetics that do not induce vasodilation on cerebral blood flow.

Our proposed method enabled repeatable longitudinal brain imaging over a three-week period, addressing a key limitation of conventional ULM imaging and offering potential for various preclinical applications. However, there are still some limitations in this study.

One of the limitations is the lack of objective measures to assess the effectiveness of head-fix habituation in reducing anxiety. This may introduce variability in stress levels among mice. Recent studies suggest that tracking physiological parameters such as heart rate, respiratory rate, and corticosterone levels during habituation can confirm that mice reach a low stress state prior to imaging48. This approach would be highly beneficial for future awake imaging studies. Furthermore, alternative head-fixation setups, such as air-floated balls or treadmills, which allow the free movement of limbs, have been shown to reduce anxiety and facilitate natural behaviors during imaging30. Adopting these approaches in future studies could enhance the reliability of awake imaging data by minimizing stress-related confounds.

Earlier version: Another limitation of this study is the potential residual vasodilatory effect of isoflurane anesthesia on awake imaging sessions. The awake imaging sessions were conducted shortly after the mice had emerged from isoflurane anesthesia, required for the MB bolus injections. The lasting vasodilatory effects of isoflurane may have influenced vascular responses, potentially contributing to an underestimation of differences in vascular dynamics between anesthetized and awake state. Future applications of awake ULM in functional imaging using an indwelling jugular vein catheter presents a promising alternative to enable more accurate functional imaging in awake animals, addressing current limitations associated with anesthesia-induced vascular effects.

Revised version: Another limitation of this study is the potential residual vasodilatory effect of isoflurane anesthesia on awake imaging sessions and the short imaging window available after bolus injection. The awake imaging sessions were conducted shortly after the mice had emerged from isoflurane anesthesia, required for the MB bolus injections. The lasting vasodilatory effects of isoflurane may have influenced vascular responses, potentially contributing to an underestimation of differences in vascular dynamics between anesthetized and awake state. In addition, since microbubbles are rapidly cleared from circulation, the duration of effective imaging is limited to only a few minutes, which also overlaps with the anesthesia recovery period, constraining the usable awake-state imaging window. Future improvement on microbubble infusion using an indwelling jugular vein catheter presents a promising alternative to address these limitations. This method allows for stable microbubble infusion without the need for anesthesia induction, ensuring that the awake imaging condition is free from residual anesthetic effects. Moreover, it has the potential to extend the duration of imaging sessions, offering a longer and more stable time window for data acquisition. Furthermore, by performing ULM imaging in the awake state first, instead of starting with anesthetized imaging, researchers can achieve a more rigorous comparison of how various anesthetics influence cerebral microvascular dynamics relative to the awake baseline.

In our longitudinal study, consistent imaging results were obtained over a three-week period, demonstrating the feasibility of awake ULM imaging for this duration. However, for certain research applications, a monitoring period of several months would be valuable. Extending the duration of longitudinal awake ULM imaging to enable such long-term studies is a potential direction for future development.

Earlier version: Tissue motion is also a critical concern of ULM imaging. Benefiting from the robust head-fixing protocols and successful habituation, we did not observe significant motion artifacts in the ULM images. Nevertheless, we still employed the correlation method to register the ULM images49 in order to remove subtle tissue movement and improve the ULM image quality. This method preserves the maximum fidelity for the MB tracks, but may not be adequate for non-rigid motion. While rigid motion correction is often effective in anesthetized animals, awake animal imaging presents greater challenges due to the more prominent non-rigid motion, particularly in deeper brain regions. This is evidenced in Supplementary Fig. 1 (Mouse 7), where cortical vessels remain relatively stable, but regions around the colliculi and mesencephalon exhibit more noticeable motion artifacts, indicating that displacement is more pronounced in deeper areas. To address these deeper, non-rigid motions, recent studies suggest estimating non-rigid transformations from unfiltered tissue signals before applying corrections to ULM vascular images16,50. Such advanced motion correction strategies may be more effective for awake ULM imaging, which experiences higher motion variability. The development of more robust and effective motion correction techniques will be crucial to reduce motion artifacts in future awake ULM applications. Meanwhile, with 2D imaging, we cannot correct for out-of-plane motion, which necessitates 3D imaging. In the future, 3D motion correction techniques that account for complex tissue motions and are computationally efficient need to be developed for awake and longitudinal ULM imaging.

Revised version: Tissue motion is also a critical concern of ULM imaging. While rigid motion correction is often effective in anesthetized animals, awake animal imaging presents greater challenges due to the more prominent non-rigid motion, particularly in deeper brain regions. This is evidenced in Supplementary Fig. 1 (Mouse 7), where cortical vessels remain relatively stable, but regions around the colliculi and mesencephalon exhibit more noticeable motion artifacts, indicating that displacement is more pronounced in deeper areas. To address these deeper, non-rigid motions, recent studies suggest estimating non-rigid transformations from unfiltered tissue signals before applying corrections to ULM vascular images16, 49. Such advanced motion correction strategies may be more effective for awake ULM imaging, which experiences higher motion variability. The development of more robust and effective motion correction techniques will be crucial to reduce motion artifacts in future awake ULM applications. Meanwhile, with 2D imaging, we cannot correct for out-of-plane motion, which necessitates 3D imaging. In the future, 3D motion correction techniques that account for complex tissue motions and are computationally efficient need to be developed for awake and longitudinal ULM imaging.

Removed in the revised version: Based on the maximum linking distance and gap closing parameters outlined in the Methods section, blood flow with velocities below 50 mm/s can be detected. However, the use of a directional filter to estimate flow direction may introduce aliasing. MBs moving at higher velocities may be subject to incorrect flow direction estimation due to aliasing effects. Given that the compounded frame rate is 1000 Hz, with an ultrasound center frequency of 20 MHz and a sound speed of 1540 m/s, the relationship between Doppler frequency and the axial blood flow velocity12 indicates that aliasing will not occur for axial flow velocities below 19.25 mm/s. In all flow velocity maps presented in this study, the range is limited to a maximum of 15 mm/s, remaining below the critical threshold for aliasing. Additionally, all vessels analyzed in the violin plots for arteriovenous flow comparisons fall within this range. While cortical arterioles and venules generally exhibit moderate flow speeds, aliasing remains a factor to consider when combining directional filtering with velocity analysis.

Earlier version: Advances in ULM imaging methods can benefit longitudinal awake imaging. For instance, dynamic ULM can differentiate between arteries and veins by leveraging pulsatility features51. 3D ULM, with volumetric imaging array52,53, enables the reconstruction of whole-brain vascular network, providing a more comprehensive understanding of vessel branching patterns. Meanwhile, 3D ULM also helps to mitigate the challenge of aligning the identical coronal plane for longitudinal imaging, a process that requires precise manual alignment in 2D ULM to ensure consistency. Additionally, this alignment issue can also be alleviated in 2D imaging using backscattering amplitude method, which may assist in estimating out-of-plane positioning during longitudinal imaging54.

Revised version: Advances in ULM imaging methods can benefit longitudinal awake imaging. For instance, dynamic ULM can differentiate between arteries and veins by leveraging pulsatility features50. 3D ULM, with volumetric imaging array51, 52, enables the reconstruction of whole-brain vascular network, providing a more comprehensive understanding of vessel branching patterns. Meanwhile, 3D ULM also helps to mitigate the challenge of aligning the identical coronal plane for longitudinal imaging, a process that requires precise manual alignment in 2D ULM to ensure consistency. Additionally, this alignment issue can also be alleviated in 2D imaging using backscattering amplitude method, which may assist in estimating out-of-plane positioning during longitudinal imaging53.

Earlier version: Longitudinal brain imaging in the awake state offers a promising tool for neuroscience research as it not only avoids the confounding effects of anesthesia on cerebral vasculature, but also enables observations of intrinsic dynamics of the vasculature within the same subject, minimizing potential sources of bias associated with inter-subject variability. In the future, this technique is expected to be further integrated with disease models to study the changes in cerebral vasculature during the development of diseases. Also, this technique can be further combined with the latest functional ULM (fULM) studies16 to allow awake fULM imaging. Our study laid the foundation for these studies with awake fULM, which is expected to improve the sensitivity of conventional fULM techniques because hemodynamic responses are much stronger in the awake state than in anesthesia24–26.

Revised version: Longitudinal brain imaging in the awake state offers a promising tool for neuroscience research as it not only avoids the confounding effects of anesthesia on cerebral vasculature, but also enables observations of intrinsic dynamics of the vasculature within the same subject, minimizing potential sources of bias associated with inter-subject variability. In the future, this technique is expected to be further integrated with disease models to study the changes in cerebral vasculature during the development of diseases. Also, this technique can be further combined with the latest functional ULM (fULM) studies16 to allow awake fULM imaging. Our study laid the foundation for these studies with awake fULM, which is expected to improve the sensitivity of conventional fULM techniques because hemodynamic responses are much stronger in the awake state than in anesthesia24–26. However, it is also important to note that although longitudinal awake imaging presents promise to avoid the confounding effects of anesthetics, imaging under anesthesia remains more convenient and controllable in many cases. For applications where the physiological question of interest is not sensitive to anesthesia-induced vascular effects, anesthetized imaging still offers a simpler and more stable approach. Awake imaging inherently exhibits greater physiological variability. However, care must be taken at the experimental level to minimize confounding sources of variation, such as stress level of the animal or handling inconsistencies, to ensure that the measurements are physiologically meaningful.

Methods

Earlier version: Animal preparation

Revised version: Animal preparation and usage overview

Earlier version: Six healthy female C57 mice (8-12 weeks) were used for this study. All experimental procedures were conducted in accordance with the guidelines set by the University of Illinois Institutional Animal Care and Use Committee (IACUC Protocol number #22165). The animals were housed in an environment with a 12-hour light/dark cycle and had free access to food and water. Prior to the cranial window surgery, the mice were kept in group housing, and they were individually housed post-surgery.

Revised version: Eight healthy female C57 mice (8-12 weeks) were used for this study, numbered as Mouse 1 to Mouse 8. Three mice (Mouse 1–3) were used to compare imaging results between awake and anesthetized states (Fig. 3 and 4). Three additional mice (Mouse 4–6) underwent longitudinal imaging over a three-week period (Fig. 5 and 6). Among them, Mouse 4 was also used as an example to demonstrate the overall system schematic and saturation conditions (Fig. 1 and 2). Several mice (Mouse 2, 6, 7, and 8) exhibited suboptimal cranial window quality or image artifacts and were included to illustrate common surgical or imaging issues (Supplementary Fig. 1). The specific usage of each animal is also annotated in the corresponding figure captions. All experimental procedures were conducted in accordance with the guidelines set by the University of Illinois Institutional Animal Care and Use Committee (IACUC Protocol number #22165). The animals were housed in an environment with a 12-hour light/dark cycle and had free access to food and water. Prior to the cranial window surgery, the mice were kept in group housing, and they were individually housed post-surgery.

Pre-surgery handling

Earlier version: A week before performing the cranial window surgery, the animals underwent tunnel handling, a procedure shown to significantly reduce the levels of anxiety in mice55,56. Specifically, a commercially available polycarbonate mouse transfer tube (TRANS-TUBE 130X50MM, Braintree Scientific, Inc.) was used for the initial three days. The mice were encouraged to enter the tube from their cage, after which the tube was lifted, allowing the mouse to remain inside for 30 seconds. The mouse was then returned to its cage for free movement for one minute, before a second identical tunnel handling was carried out. This procedure was done twice daily. Approximately three days into this routine, the mice were accustomed to the tunnel handling method. The animals were then picked up using the commercial tunnel, and then a 3D-printed body tube was attached to the other end of the tunnel. The mice were allowed to enter the body tube voluntarily. Subsequently, the body tube was utilized for handling the mice in the following 5 days as a replacement for the commercial tunnel.

Revised version: A week before performing the cranial window surgery, the animals underwent tunnel handling, a procedure shown to significantly reduce the levels of anxiety in mice54, 55. Specifically, a commercially available polycarbonate mouse transfer tube (TRANS-TUBE 130X50MM, Braintree Scientific, Inc.) was used for the initial three days. The mice were encouraged to enter the tube from their cage, after which the tube was lifted, allowing the mouse to remain inside for 30 seconds. The mouse was then returned to its cage for free movement for one minute, before a second identical tunnel handling was carried out. This procedure was done twice daily. Approximately three days into this routine, the mice were accustomed to the tunnel handling method. The animals were then picked up using the commercial tunnel, and then a 3D-printed body tube was attached to the other end of the tunnel. The mice were allowed to enter the body tube voluntarily. Subsequently, the body tube was utilized for handling the mice in the following 5 days as a replacement for the commercial tunnel.

Cranial window surgery

The entire procedure closely followed a previously published protocol33. To minimize brain swelling, the animal received an intraperitoneal injection of Dexamethasone (0.5 mg/kg body weight). The mouse was anesthetized by inhalation of isoflurane (3% for induction and 1%-1.5% for maintenance). After confirming that the mouse was anesthetized, the head of the animal was fixed in the stereotaxic frame. Afterward, the scalp of the mouse was incised, and the temporalis muscle was dissected from the skull using forceps. Tissue adhesive (3M, cat. no. 70200742529) was applied to secure the retracted muscles and incision edges, ensuring proper closure of the skin. Subsequently, a headpost was fixed to the skull using tissue adhesive and dental cement (Sun Medical Co., Ltd, cat. no. Super-Bond C&B). The skull was then carefully thinned using a high-speed rotary micromotor (Foredom, cat. no. K.1070) to create a cranial window with a width of approximately 8 mm laterally and which extending from bregma to lambda in the anterior-posterior direction. The skull was continuously thinned along this outline until it became loose, allowing for separation of the bone from the dura mater using forceps. After removing the skull, a polymethyl pentene plastic sheet (Goodfellow, cat. no. ME311051) was placed over the cranial window and fixed to the bone using tissue adhesive and dental cement. To provide protection, the window was covered with biocompatible silicone rubber (Smooth-on, cat. no. Body Double-Fast Set). Once the silicone rubber was secure, the anesthesia was discontinued, and the mouse was allowed to regain consciousness.

Post-surgical care and head-fix habituation

Following the surgery, subcutaneous administration of Carprofen (5-10 mg/kg body weight) was provided to the animals for immediate post-operative pain management. The animals were allowed to recover from anesthesia in their individual cages, with a heating pad provided to maintain optimal body temperature during the recovery period. Observations were made every 15 minutes until the animals reached sternal recumbency. To aid in their recovery, dry food soaked in water was provided, along with the use of recovery gel (ClearH2O, cat. no. 72-06-5022) to facilitate chewing and hydration. In the case that any signs of pain were observed, additional subcutaneous doses of 5-10 mg/kg Carprofen were administered every 12 to 24 hours to alleviate surgical discomfort for 3 days. Post-operative monitoring of the animals was performed daily for a duration of 14 days following the surgery. Once the animals had fully recovered from the surgical procedure, the animals were habituated daily to head fixation after walking through the body tube. The duration of the head-fixation periods was gradually increased over time, starting at approximately 10 minutes on day 1 and extending to up to 1 hour after 3 days33. If any signs of discomfort, such as excessive movement or vocalization, were observed during the head-fixation sessions then the sessions would be immediately discontinued.

The use of 3D-printed body tube

The body tube was secured to the table using four screws on the outer side, while two small screws on the top were used to clamp the headpost. Two semicircular grooves that fit the size of the small screws were made on the headpost (Fig. 1a). During the process of head-fixing the animal, the experimenter guided the animal to enter the body tube from the rear. Once the animal protruded its head from the front end, the experimenter manually grasped the headpost, gently restrained the head of the animal, and subsequently aligned the semicircular grooves with the screws to immobilize the animal.

Experimental procedure of imaging sessions

On the day of imaging, the animal was guided to walk into the body tube, and the headpost was firmly secured to the tube. The protective silicone rubber and headpost were cleaned using 70% ethanol, followed by rinsing with sterile saline. Then, forceps were inserted between the silicone rubber and the cement to detach the silicone rubber from the headpost. Due to the challenges of conducting continuous infusion in awake animals, contrast microbubbles were administered via tail vein bolus injection in this study. Animals were anesthetized with isoflurane (3% for induction and 1% for maintenance) before tail vein catheterization to alleviate pain and stress. Ultrasound coupling gel was applied, and the transducer was positioned to find the imaging plane. When the imaging plane was identified, the relative position of the probe and the headpost was recorded, serving as a reference point for subsequent longitudinal studies. B-mode and power Doppler images were also saved as references to facilitate the identification of the same imaging plane in the following weeks.

Once the imaging plane was confirmed, the isoflurane anesthesia was terminated, and 0.1 ml of DEFINITY® (Lantheus, North Billerica, MA) microbubble was administered via the tail vein catheter. The completion of the injection was considered as the starting point, denoted as T = 0. From T = 0 onwards, continuous ultrasound data was saved. After completion of the imaging session, the cranial window was filled with the same biocompatible silicone rubber described above, allowing approximately 10 minutes for solidification. The headpost was then removed from the body tube, and the animal was returned to its home cage. Imaging sessions were conducted once per week over a three-week period.

For the experiment of comparing differences between anesthesia and awake state (Fig. 3-5), once the mice were fixed to the body tube, they were continuously anesthetized with 2% isoflurane for more than 15 minutes. During this process, the tail vein catheterization was completed, and the imaging plane was confirmed. Following the bolus injection of 0.1 ml MB, ULM data under anesthesia was acquired for 1000 seconds. The tail vein catheter was remained in place. After the ULM acquisition under anesthesia, it was essential to wait until no bubbles were circulating within the blood stream before terminating anesthesia. Once the animal approached full arousal, another bolus injection of 0.1 ml MB was administered. The catheter was then removed, and ULM acquisition started (marked as T = 0). Other procedures were the same as those described in the previous two paragraphs. For the ULM data collected under anesthesia, since there was no specific time point at which mice achieved full pupillary dilatation, we consistently employed the RF dataset starting from T = 500s until the cumulative MB count reached 5 million for ULM reconstruction to make the MB concentration similar as in the awake cases.

Pupillary recording and measurement

Earlier version: Monitoring the pupillary area of rodents is commonly used to determine their level of arousal57. The pupils constrict under isoflurane anesthesia and enlarge upon awakening38,58. In this study, video recording of pupillary area performed using a camera (ace acA800-510um, Basler Inc., Exton, PA) placed in front of the mouse, with timestamps synchronized to the ultrasound acquisition. The ULM imaging session was carried out in a dark room. An infrared (IR) flashlight (EVOLVA Future Technology T20) was positioned above the camera. Since mice are functionally blind to IR wavelengths59,60, the use of IR illumination can minimize the effect of light exposure on pupil changes so that the pupillary area is mostly influenced by anesthesia. The videos recorded by the camera were analyzed using ImageJ software, and the contour of the pupil was manually circled. The equivalent area of a single pixel was calibrated, which allowed for the quantification of the pupillary area according to the number of pixels occupied by the pupil.

Revised version: Monitoring the pupillary area of rodents is commonly used to determine their level of arousal56. The pupils constrict under isoflurane anesthesia and enlarge upon awakening38, 57. In this study, video recording of pupillary area performed using a camera (ace acA800-510um, Basler Inc., Exton, PA) placed in front of the mouse, with timestamps synchronized to the ultrasound acquisition. The ULM imaging session was carried out in a dark room. An infrared (IR) flashlight (EVOLVA Future Technology T20) was positioned above the camera. Since mice are functionally blind to IR wavelengths58, 59, the use of IR illumination can minimize the effect of light exposure on pupil changes so that the pupillary area is mostly influenced by anesthesia. The videos recorded by the camera were analyzed using ImageJ software, and the contour of the pupil was manually circled. The equivalent area of a single pixel was calibrated, which allowed for the quantification of the pupillary area according to the number of pixels occupied by the pupil.

Ultrasound imaging sequence

Earlier version: A Vantage 256 system (Verasonics Inc., Kirkland, WA) was used for this experiment. The ultrasound system was connected to a linear-array transducer (L35-16vX, Verasonics Inc.), which was fixed onto a 3D-printed holder attached to a translation motor (VT-80, Physik Instrumente, Auburn, MA). The motor allowed precise control of transducer movement in the elevational direction, enabling adjustment of the imaged coronal plane. Ultrasound was transmitted at a center frequency of 20 MHz. A 9-angle compounded plane wave technique was used (angles: -8°, -6°, -4°, -2°, 0°, 2°, 4°, 6°, and 8°), with a post-compounded frame rate of 1000 Hz. Interleaved sampling is employed to capture high-frequency echoes more effectively. With the system’s sampling rate limited to 62.5 MHz, the upper limit of the center frequency of the transducer passband is 15.625 MHz. To mitigate aliasing, two transmissions are sent per angle, staggered in time. This approach effectively doubles the sampling rate, ensuring more accurate image reconstruction. The ensemble size of each dataset was 800 frames, and the acquired radiofrequency (RF) data was saved for ofline reconstruction. Beamforming was conducted using the Verasonics built-in program to reconstruct the in-phase and quadrature (IQ) data. The IQ data had a pixel size of half the wavelength in the axial direction and one wavelength in the lateral direction.

Revised version: A Vantage 256 system (Verasonics Inc., Kirkland, WA) was used for this experiment. The ultrasound system was connected to a linear-array transducer (L35-16vX, Verasonics Inc.), which was fixed onto a 3D-printed holder attached to a translation motor (VT-80, Physik Instrumente, Auburn, MA). The motor allowed precise control of transducer movement in the elevational direction, enabling adjustment of the imaged coronal plane. Ultrasound was transmitted at a center frequency of 20 MHz. A 9-angle compounded plane wave technique was used (angles: -8°, -6°, -4°, -2°, 0°, 2°, 4°, 6°, and 8°), with a post-compounded frame rate of 1000 Hz. Interleaved sampling is employed to capture high-frequency echoes more effectively. With the system’s sampling rate limited to 62.5 MHz, the upper limit of the center frequency of the transducer passband is 15.625 MHz. To mitigate aliasing, two transmissions are sent per angle, staggered in time. This approach effectively doubles the sampling rate, ensuring more accurate image reconstruction. The ensemble size of each dataset was 800 frames, and the acquired radiofrequency (RF) data was saved for offline reconstruction. Beamforming was conducted using the Verasonics built-in program to reconstruct the in-phase and quadrature (IQ) data. The IQ data had a pixel size of half the wavelength in the axial direction and one wavelength in the lateral direction.

ULM image processing

Earlier version: To ensure consistent localization of bubbles at different depths, a noise equalization method was employed to adjust signal intensity61. High pass filtering with a cutoff frequency of 30Hz was applied to the IQ data to enhance sensitivity for MB extraction. MBs moving toward and away from the transducer were separated into two distinct datasets (upward flow and downward flow) based on the positive and negative Doppler shifts62. Subsequently, singular value decomposition (SVD) filtering was applied to the IQ data to further eliminate clutter signals from tissue and extract MB signals63,64. The singular value cutoff was adaptively determined to achieve objective and stable filtering results65. The filtered MB data was then spline interpolated to a resolution with a lateral and axial pixel size of one-tenth of the wavelength (4.928 µm). Each interpolated frame was subjected to 2D normalized cross-correlation with the point spread function of the imaging system, which was empirically determined. Regional maxima in the cross-correlation results indicated the center position of MBs49. Thresholding of image intensity was applied to reject low-intensity values and prevent noise from being erroneously identified as bubbles. The coordinates of the microbubble center points obtained from each frame were tracked using the uTrack algorithm66,67. To ensure the detection of reliable MB tracks, only tracks with a minimum length of 10 consecutive frames were considered valid. The choice of 10 consecutive frames (10 ms) was based on established practice11 but can be adjusted as needed68. For the uTrack algorithm, two additional key parameters were specified: the maximum linking distance and the gap-filling distance, both set to 10 pixels (∼50 microns). This configuration means that only bubble centroids within 10 pixels of each other across consecutive frames are considered part of the same bubble trajectory. Additionally, when the start and end points of two tracks fall within this threshold, the gap-filling parameter merges them into a single, continuous track. It is important to select these parameters carefully, as overly large values could lead to an overestimation of flow velocity. By setting the maximum linking distance to 10 pixels, we effectively limited the measurable velocity to 50 mm/s, under the assumption that no bubble would exceed a 50-micron displacement within the 1 ms interval between frames. Considering the velocity distribution across the mouse brain69, this 50 mm/s limit ensures that most of the blood flow is captured accurately. After determining bubble tracks with the specified parameters for uTrack algorithm, accumulating the MB tracks resulted in the flow intensity map. The final values were square-root transformed to compress the dynamic range of the image for display. Furthermore, the distance traveled by MB between adjacent frames was calculated during the tracking process to determine blood flow velocity, which was then assigned to each individual bubble track. The average velocity computed from multiple tracks was used to generate a comprehensive flow speed map.

Revised version: To ensure consistent localization of bubbles at different depths, a noise equalization method was employed to adjust signal intensity60. High pass filtering with a cutoff frequency of 30Hz was applied to the IQ data to enhance sensitivity for MB extraction. MBs moving toward and away from the transducer were separated into two distinct datasets (upward flow and downward flow) based on the positive and negative Doppler shifts61. Subsequently, singular value decomposition (SVD) filtering was applied to the IQ data to further eliminate clutter signals from tissue and extract MB signals62, 63. The singular value cutoff was adaptively determined to achieve objective and stable filtering results64. The filtered MB data was then spline interpolated to a resolution with a lateral and axial pixel size of one-tenth of the wavelength (4.928 µm). Each interpolated frame was subjected to 2D normalized cross-correlation with the point spread function of the imaging system, which was empirically determined. Regional maxima in the cross-correlation results indicated the center position of MBs65. Thresholding of image intensity was applied to reject low-intensity values and prevent noise from being erroneously identified as bubbles. The coordinates of the microbubble center points obtained from each frame were tracked using the uTrack algorithm66, 67. To ensure the detection of reliable MB tracks, only tracks with a minimum length of 10 consecutive frames were considered valid. The choice of 10 consecutive frames (10 ms) was based on established practice11 but can be adjusted as needed68. For the uTrack algorithm, two additional key parameters were specified: the maximum linking distance and the gap-filling distance, both set to 10 pixels (∼50 microns). This configuration means that only bubble centroids within 10 pixels of each other across consecutive frames are considered part of the same bubble trajectory. Additionally, when the start and end points of two tracks fall within this threshold, the gap-filling parameter merges them into a single, continuous track. It is important to select these parameters carefully, as overly large values could lead to an overestimation of flow velocity. By setting the maximum linking distance to 10 pixels, we effectively limited the measurable velocity to 50 mm/s, under the assumption that no bubble would exceed a 50-micron displacement within the 1 ms interval between frames. Considering the velocity distribution across the mouse brain69, this 50 mm/s limit ensures that most of the blood flow is captured accurately. After determining bubble tracks with the specified parameters for uTrack algorithm, accumulating the MB tracks resulted in the flow intensity map. The final values were square-root transformed to compress the dynamic range of the image for display. Furthermore, the distance traveled by MB between adjacent frames was calculated during the tracking process to determine blood flow velocity, which was then assigned to each individual bubble track. The average velocity computed from multiple tracks was used to generate a comprehensive flow speed map.

Quantitative analysis of ULM images

For single-vessel analysis, the vessel density map (either bidirectional or unidirectional) was binarized, and the center line of the vessel was extracted using MATLAB “bwskel” function. The radius of each vessel was defined as the distance from the central line to the nearest vessel edge on the binarized map, with diameter calculated by doubling this radius. Given a pixel size of approximately 4.928 µm, the minimum detectable vessel diameter using this approach is 9.856 µm, explaining why the smallest measured radius in unsaturated conditions, as shown in Fig. 2b, always falls into this value. To analyze a specific vessel, we manually selected a vessel segment. Each point on the central line provided a diameter measurement. For flow velocity, each non-zero pixel within the chosen segment was used to estimate effective blood flow velocity in the vessel.

Earlier version: For ROI-based analysis, ROIs were manually defined based on a brain atlas to ensure consistency across both anesthetized and awake states, as well as across weekly imaging sessions, enhancing comparability of results. In ROI-based analysis, we focused on two primary parameters: vascularity and mean velocity. Vascularity was defined as the proportion of the pixel count occupied by blood vessels within each ROI, obtained by binarizing the ULM vessel density maps and calculating the percentage of the pixels with MB signal. Mean velocity was calculated by averaging all non-zero pixel of velocity estimates within the ROI. The velocity distribution within each ROI was also visualized using violin plots, as shown in Fig. 2, 4 and 6, to illustrate the range and density of flow velocity estimates across different acquisition.

Revised version: For ROI-based analysis, ROIs were manually defined based on a brain atlas to ensure consistency across both anesthetized and awake states, as well as across weekly imaging sessions, enhancing comparability of results. In ROI-based analysis, we focused on two primary parameters: fractional vessel area and mean velocity. Fractional vessel area was defined as the proportion of the pixel count occupied by blood vessels within each ROI, obtained by binarizing the ULM vessel density maps and calculating the percentage of the pixels with MB signal. Mean velocity was calculated by averaging all non-zero pixel of velocity estimates within the ROI. The velocity distribution within each ROI was also visualized using violin plots, as shown in Fig. 2, 4 and 6, to illustrate the range and density of flow velocity estimates across different acquisition. In this study, we focused on these two metrics because they represent the most straightforward extension of single-vessel analysis to brain-wide vascular changes.